<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Activities: Week 3 &#8212; Programs, Information, and People: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/matrixeq.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webgldemo.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webglinteractive.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/js_lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../_static/accessibility.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="https://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Activities: Week 4" href="week04.html" />
    <link rel="prev" title="Activities: Week 2 (Note that PS1 is found in Week 1)" href="week02.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'pip2';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = true;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Table of Contents</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div class="gittip">
                  <iframe style="border:0;"
                    src="https://www.gittip.com/bnmnetp/widget.html"
                    width="48pt" height="22pt" class="gittip-button">
                  </iframe>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/pip2/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
<!--            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li> -->
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
<!--            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> --><!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/pip2/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='https://github.com/bnmnetp/runestone/issues/new'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Activities: Week 3</a><ul>
<li><a class="reference internal" href="#problem-set">Problem Set</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Activities: Week 2 (Note that PS1 is found in Week 1)" data-toggle="tooltip">
    <a href="week02.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Activities: Week 4' data-toggle="tooltip" >
    <a href="week04.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="activities-week-3">
<h1>Activities: Week 3<a class="headerlink" href="#activities-week-3" title="Permalink to this headline">¶</a></h1>
<p>You have the following activities to complete:</p>
<ul class="simple">
<li><strong>Before class Tuesday:</strong><ul>
<li>Read <a class="reference internal" href="../Iteration/intro-Iteration.html#iteration-chap"><span class="std std-ref">Iteration</span></a>.</li>
<li>And try as many of the exercises as you can, thinking about them and noting what’s confusing, what’s hard</li>
</ul>
</li>
<li><strong>Before class Thursday:</strong><ul>
<li>Read <a class="reference internal" href="../Selection/ConditionalExecutionBinarySelection.html#conditionals-chap"><span class="std std-ref">Conditionals</span></a></li>
<li>Read <a class="reference internal" href="../Files/intro-WorkingwithDataFiles.html#files-chap"><span class="std std-ref">File Input/Output</span></a></li>
<li>Additionally, here are some more introductory videos about concepts we are looking at this week. You may find them very useful; I recommend you watch/listen to them:<ul>
<li><a class="reference external" href="https://www.youtube.com/watch?v=3GsNQkIskaY">Split and Join string methods</a></li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=HVoutOx3hMM">Code blocks</a></li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=gIlOfoy4iZk">For loops and iteration</a></li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=nT7q2JTNHog">The accumulation pattern</a></li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=YUAUWU9Pank">Booleans</a></li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=V3K44gSruoY">Conditionals</a></li>
</ul>
</li>
</ul>
</li>
<li>By <strong>Sunday at 11:59PM</strong>, save answers to the exercises in <strong>Problem Set 2</strong>:<ul>
<li>Complete each of the problem set problems.</li>
<li>Submit your Demonstrate Your Understanding 3 assignment.</li>
<li>Resubmit your last DYU (DYU 2) if you did not get full credit on it.</li>
</ul>
</li>
</ul>
<div class="section" id="problem-set">
<span id="problem-set-2"></span><h2>Problem Set<a class="headerlink" href="#problem-set" title="Permalink to this headline">¶</a></h2>
<p><strong>Instructions:</strong> Write the code you want to save in the provided boxes, and click <strong>save &amp; run</strong> for each one. The last code you have saved for each one by the deadline is what will be graded.</p>

<pre data-component="datafile" id=about_programming.txt data-hidden data-edit="false" data-rows="20" data-cols="65">
Computer programming (often shortened to programming) is a process that leads from an
original formulation of a computing problem to executable programs. It involves
activities such as analysis, understanding, and generically solving such problems
resulting in an algorithm, verification of requirements of the algorithm including its
correctness and its resource consumption, implementation (or coding) of the algorithm in
a target programming language, testing, debugging, and maintaining the source code,
implementation of the build system and management of derived artefacts such as machine
code of computer programs. The algorithm is often only represented in human-parseable
form and reasoned about using logic. Source code is written in one or more programming
languages (such as C++, C#, Java, Python, Smalltalk, JavaScript, etc.). The purpose of
programming is to find a sequence of instructions that will automate performing a
specific task or solve a given problem. The process of programming thus often requires
expertise in many different subjects, including knowledge of the application domain,
specialized algorithms and formal logic.
Within software engineering, programming (the implementation) is regarded as one phase in a software development process. There is an on-going debate on the extent to which
the writing of programs is an art form, a craft, or an engineering discipline. In
general, good programming is considered to be the measured application of all three,
with the goal of producing an efficient and evolvable software solution (the criteria
for "efficient" and "evolvable" vary considerably). The discipline differs from many
other technical professions in that programmers, in general, do not need to be licensed
or pass any standardized (or governmentally regulated) certification tests in order to
call themselves "programmers" or even "software engineers." Because the discipline
covers many areas, which may or may not include critical applications, it is debatable
whether licensing is required for the profession as a whole. In most cases, the
discipline is self-governed by the entities which require the programming, and sometimes
very strict environments are defined (e.g. United States Air Force use of AdaCore and
security clearance). However, representing oneself as a "professional software engineer"
without a license from an accredited institution is illegal in many parts of the world.
</pre>

<div data-childcomponent="ps_1_04" class="runestone explainer ac_section alert alert-warning">
Write one for loop to print out each element of the list <code class="docutils literal notranslate"><span class="pre">several_things</span></code>. Then, write <em>another</em> for loop to print out the TYPE of each element of the list called <code class="docutils literal notranslate"><span class="pre">several_things</span></code>. To complete this problem you should have written two different for loops, each of which iterates over the list <code class="docutils literal notranslate"><span class="pre">several_things</span></code>, but each of those 2 for loops should have a different result.
<textarea data-component="activecode" id=ps_1_04 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
several_things = ["hello", 2, 4, 6.0, 7.5, 234352354, "the end", "", 99]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
      self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
      self.assertIn("<class 'str'>\n<class 'int'>\n<class 'int'>\n<class 'float'>\n<class 'float'>\n<class 'int'>\n<class 'str'>\n<class 'str'>\n<class 'int'>", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_1_06" class="runestone explainer ac_section alert alert-warning">
<p>Write code that uses iteration to print out each element of the list stored in <code class="docutils literal notranslate"><span class="pre">excited_words</span></code>, BUT print out each element <strong>without</strong> its ending punctuation. You should see:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">hello</span>
<span class="n">goodbye</span>
<span class="n">wonderful</span>
<span class="n">I</span> <span class="n">love</span> <span class="n">Python</span>
</pre></div>
</div>
<p>(Hint: remember string slicing?)</p>

<textarea data-component="activecode" id=ps_1_06 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
excited_words = ["hello!", "goodbye!", "wonderful!", "I love Python?"]

# Write your code here.
=====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def test_output(self):
        self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
        self.assertIn("hello\ngoodbye\nwonderful\nI love Python", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_2_01" class="runestone explainer ac_section alert alert-warning">
Write code to count the number of characters in <code class="docutils literal notranslate"><span class="pre">original_str</span></code> using the accumulation pattern and assign the answer to a variable <code class="docutils literal notranslate"><span class="pre">num_chars_sent</span></code>. Do NOT use the <code class="docutils literal notranslate"><span class="pre">len</span></code> function to solve the problem (if you use it while you are working on this problem, comment it out afterward!)
<textarea data-component="activecode" id=ps_2_01 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
original_str = "The quick brown rhino jumped over the extremely lazy fox."


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(num_chars_sent, len(original_str), "Testing whether num_chars_sent has the correct value")
       self.assertNotIn('len', self.getEditorText(), "Testing that you are not including the len function in your code. (Don't worry about Actual and Expected Values.)")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_2_02" class="runestone explainer ac_section alert alert-warning">
Write code to open the file <code class="docutils literal notranslate"><span class="pre">about_programming.txt</span></code> which has been provided for you in this problem set, and assign the <strong>number of lines</strong> in the file to the variable <code class="docutils literal notranslate"><span class="pre">file_lines_num</span></code>.
<textarea data-component="activecode" id=ps_2_02 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio='about_programming.txt*atype*'   
          >
# Write your code here.

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

 def testOne(self):
    self.assertIn('open', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
    self.assertEqual(file_lines_num,len(open("about_programming.txt","r").readlines()), "Testing to see that file_lines_num has been set to the number of lines in the file.")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_2_03" class="runestone explainer ac_section alert alert-warning">
<p>The program below doesn’t always work as intended. Try uncommenting different lines setting the initial value of x. Tests will run at the end of your code, and you will get diagnostic error messages.</p>
<p>Fix the code so that it passes the test for each different value of x. So when the first line is uncommented, and when the second line, third line, and fourth line are each uncommented, you should always pass the test.</p>
<p>(HINT: you don’t have to make a big change!)</p>

<textarea data-component="activecode" id=ps_2_03 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
#x = 25
#x = 15
#x = 5
#x = -10

if x > 20:
  y = "oranges"
if x > 10:
  y = "strawberries"
if x < 0:
  y = "blueberries"
else:
  y = "mulberries"

print("y is " + str(y))

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      print("No tests for the comment, of course -- we can only test stored values!\n")
      if x == 25:
          self.assertEqual(y, "oranges", "test when x is 25: y should be 'oranges'")
      elif x == 15:
          self.assertEqual(y, 'strawberries', "test when x is 15: y should be 'strawberries'")
      elif x == 5:
          self.assertTrue(y=='mulberries' or y=="unknown", "test when x is 5: y should be 'mulberries' (but if you passed this test at any point, even when we had an error in the PSet, it will still pass)")
      elif x == -10:
          self.assertEqual(y, 'blueberries', "test when x is -10: y should be 'blueberries'")
      else:
          print("No tests when value of x is %s" % (x))

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_2_04" class="runestone explainer ac_section alert alert-warning">
<p>How many characters are in each element of list <code class="docutils literal notranslate"><span class="pre">lp</span></code>? Write code to print the length (number of characters) of each element of the list, on a separate line. (Do not write 8+ lines of code to do this. Use a for loop.)</p>
<p>The output you get should be:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="mi">5</span>
<span class="mi">13</span>
<span class="mi">11</span>
<span class="mi">12</span>
<span class="mi">3</span>
<span class="mi">12</span>
<span class="mi">11</span>
<span class="mi">6</span>
</pre></div>
</div>
<p>Use iteration (a for loop).</p>

<textarea data-component="activecode" id=ps_2_04 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
lp = ["hello","arachnophobia","lamplighter","inspirations","ice","amalgamation","programming","Python"]
====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
      self.assertIn('for', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
  def test_outputA(self):
      self.assertIn("5\n13\n11\n12\n3\n12\n11\n6", self.getOutput(), "Testing output (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_2_05" class="runestone explainer ac_section alert alert-warning">
<p>Write code to count the number of strings in list <code class="docutils literal notranslate"><span class="pre">items</span></code> that have the character <code class="docutils literal notranslate"><span class="pre">w</span></code> in it. Assign that number to the variable <code class="docutils literal notranslate"><span class="pre">acc_num</span></code>.</p>
<p>HINT 1: Use the accumulation pattern!</p>
<p>HINT 2: the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator checks whether a substring is present in a string.</p>

<textarea data-component="activecode" id=ps_2_05 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
items = ["whirring", "calendar", "wry", "glass", "", "llama","tumultuous","owing"]
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertIn(' in ', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")
      self.assertEqual(acc_num, 3, "Testing that acc_num has been set to the number of strings that have 'w' in them.")

myTests().main()
</textarea>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Activities: Week 2 (Note that PS1 is found in Week 1)' data-toggle="tooltip">
    <a href="week02.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="week02.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Activities: Week 4' data-toggle="tooltip" >
    <a href="week04.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="week04.html" >Next Section - Activities: Week 4</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span> | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2013 Paul Resnick, Brad Miller, David Ranum, Created using Runestone Interactive.
    </p>
  </div>
</footer>


<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>