<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Activities: Week 1 &#8212; Programs, Information, and People: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/matrixeq.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webgldemo.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webglinteractive.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/js_lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../_static/accessibility.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="https://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Activities: Week 2 (Note that PS1 is found in Week 1)" href="week02.html" />
    <link rel="prev" title="Class Assignments" href="../assignments.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'pip2';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = true;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Table of Contents</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div class="gittip">
                  <iframe style="border:0;"
                    src="https://www.gittip.com/bnmnetp/widget.html"
                    width="48pt" height="22pt" class="gittip-button">
                  </iframe>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/pip2/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
<!--            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li> -->
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
<!--            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> --><!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/pip2/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='https://github.com/bnmnetp/runestone/issues/new'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Activities: Week 1</a><ul>
<li><a class="reference internal" href="#problem-set">Problem Set</a></li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Class Assignments" data-toggle="tooltip">
    <a href="../assignments.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Activities: Week 2 (Note that PS1 is found in Week 1)' data-toggle="tooltip" >
    <a href="week02.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="activities-week-1">
<h1>Activities: Week 1<a class="headerlink" href="#activities-week-1" title="Permalink to this headline">¶</a></h1>
<ul class="simple">
<li><strong>As soon as you can / after class 9/5:</strong><ul>
<li>Read <a class="reference external" href="https://docs.google.com/a/umich.edu/document/d/12et1AjP_on8yDwWKivSUD0SU9wAx1tvSC7y26HxcuiM/edit?usp=sharing">the course syllabus</a></li>
<li>Fill out <a class="reference external" href="https://umich.instructure.com/courses/172984/assignments/329370">this google form assignment</a></li>
<li>Check out our Piazza forums on the Canvas site.</li>
<li>Read <a class="reference external" href="http://marieflanagan.com/coding-magic/">this piece, called Coding Magic</a> by Marie LeBlanc Flanagan</li>
<li>Watch <a class="reference external" href="https://www.youtube.com/watch?v=rPAg44k6LxA">this introductory video</a>. You’ll notice it says “SI 106”, which is an undergraduate course similar to this one. That’s okay – pretend it says “506” instead of “106”. (We’ll provide more from this video series as the semester goes on, for which the same goes – pretend 106 is 506, unless of course it’s a number we’re doing math on!)</li>
</ul>
</li>
<li><strong>Before class 9/7:</strong><ul>
<li>Read <a class="reference internal" href="../GeneralIntro/intro-TheWayoftheProgram.html#the-way-of-the-program"><span class="std std-ref">General Intro</span></a>, <a class="reference internal" href="../SimplePythonData/intro-VariablesExpressionsandStatements.html#simple-python-data"><span class="std std-ref">Simple Python Data</span></a>, and <a class="reference internal" href="../PythonTurtle/intro-HelloLittleTurtles.html#turtles-chap"><span class="std std-ref">Objects and Turtle Graphics</span></a>.</li>
<li>Read <a class="reference internal" href="../Debugging/intro-HowtobeaSuccessfulProgrammer.html#debugging-chap"><span class="std std-ref">Debugging tips</span></a> and <a class="reference internal" href="../Sequences/intro-Sequences.html#sequences-chap"><span class="std std-ref">Sequences</span></a>.</li>
<li>There are also some video resources for a bunch of these topics. If you find these helpful, I strongly suggest watching them. Each is a few minutes long:<ul>
<li><a class="reference external" href="https://www.youtube.com/watch?v=e-FSJQ02bQk">Expressions</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=KcWQG6geZ_0">Variables</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=54hFgo1OAQA">Data Types</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=j_xlOtEDl_c">Print and Input</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=PdkAvKKptb8">Comments in programming</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=Exm2PL7jks8">Sequences</a>,</li>
<li><a class="reference external" href="https://www.youtube.com/watch?v=F8pL4qoaTYU">Errors</a></li>
</ul>
</li>
<li>When reading, run the activecodes and answer the multiple choice questions. There may also be exercises at the end of each chapter. These are optional (though trying them is recommended, trying all of those is not required, and getting them correct is also <em>not required</em>).</li>
<li>We are starting out with a <em>lot</em> of material. That is because we are going to speed through a lot of the beginning foundational material, and slow down in a couple weeks – this stuff is the basis of everything else we’ll do, so it is important to absorb and review as we go – you can always come back to this! When you have questions about this stuff, you should ask them, come to office hours, post on Piazza!</li>
</ul>
</li>
<li>Go to your discussion section this week, Wednesday or Thursday afternoon</li>
<li><strong>By Sunday 9/10 at 11:59 PM:</strong><ul>
<li>Complete the first <a class="reference external" href="https://umich.instructure.com/courses/172984/assignments/329347">Demonstrate Your Understanding</a> that goes with the readings.</li>
</ul>
</li>
<li><strong>By Sunday 9/17 at 11:59 PM:</strong><ul>
<li>Save answers to the exercises in <a class="reference internal" href="#problem-set-1"><span class="std std-ref">Problem Set 1</span></a>.</li>
</ul>
</li>
</ul>
<div class="section" id="problem-set">
<span id="problem-set-1"></span><h2>Problem Set<a class="headerlink" href="#problem-set" title="Permalink to this headline">¶</a></h2>
<p><strong>Instructions:</strong> Write the code you want to save in the provided boxes, and click <strong>Save &amp; Run</strong> for each one. That will  <em>run</em> your code, so you can see the output, if any, and the result of the tests, if there are any. It will also <em>save</em> your code. You should run your code each time you want to save it. Even if you only add comments, you should <strong>always click Save &amp; Run</strong>. You can then load the history of the code you have run and saved. The <em>last</em> new code you have saved for each problem by the deadline is what will be graded.</p>

<div data-childcomponent="ps_1_01" class="runestone explainer ac_section alert alert-warning">
Write code to assign the number of characters in the string <code class="docutils literal notranslate"><span class="pre">rv</span></code> to a variable <code class="docutils literal notranslate"><span class="pre">num_chars</span></code>. Then write code to assign the number of words in the string <code class="docutils literal notranslate"><span class="pre">rv</span></code> to the variable <code class="docutils literal notranslate"><span class="pre">num_words</span></code>. (Hint: remember how to split strings?)
<textarea data-component="activecode" id=ps_1_01 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
rv = """Once upon a midnight dreary, while I pondered, weak and weary,
    Over many a quaint and curious volume of forgotten lore,
    While I nodded, nearly napping, suddenly there came a tapping,
    As of some one gently rapping, rapping at my chamber door.
    'Tis some visitor, I muttered, tapping at my chamber door;
    Only this and nothing more."""

# Write your code here!

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(num_chars, len(rv), "Testing that num_chars has been set to the length of rv")
       self.assertEqual(num_words, len(rv.split()), "Testing that num_words has been set to the number of words in rv")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_1_02" class="runestone explainer ac_section alert alert-warning">
There is a function we are providing in for you in this problem set called <code class="docutils literal notranslate"><span class="pre">square</span></code>. It takes one integer and returns the square of that integer value. Write code to assign a variable called <code class="docutils literal notranslate"><span class="pre">xyz</span></code> the value <code class="docutils literal notranslate"><span class="pre">5*5</span></code> (five squared). Use the square function, rather than just multiplying with <code class="docutils literal notranslate"><span class="pre">*</span></code>.
<textarea data-component="activecode" id=ps_1_02 data-lang="python" 
     data-include="addl_functions" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
xyz = ""

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(type(xyz), type(3), "Checking type of xyz")
    def testTwo(self):
        self.assertEqual(xyz, 25, "Checking if xyz is 25")
    def testThree(self):
        self.assertIn('square', self.getEditorText(), "Testing that 'square' is in your code. (Don't worry about Actual and Expected Values.)")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_1_03" class="runestone explainer ac_section alert alert-warning">
Write in a comment next to each line of code, what each line of this code does. (You should be very specific! This exercise will train your brain for when you write more complicated code.)
<textarea data-component="activecode" id=ps_1_03 data-lang="python" 
     data-include="addl_functions" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
# Here's an example.
xyz = 12 # The variable xyz is being assigned the value 12, which is an integer

# Now do the same for each of these lines!
a = 6

b = a

# make sure to be very clear and detailed about the following line of code
orange = square(b)

print(a)

print(b)

print(orange)

pear = square

print(pear)
</textarea>
</div>

<div data-childcomponent="ps_1_05" class="runestone explainer ac_section alert alert-warning">
<p>There are a couple functions we’re giving you in this problem set. One is a function called <code class="docutils literal notranslate"><span class="pre">greeting</span></code>, which takes any string and adds <code class="docutils literal notranslate"><span class="pre">&quot;Hello,</span> <span class="pre">&quot;</span></code> in front of it. (You can see examples in the code.) Another one is a function called <code class="docutils literal notranslate"><span class="pre">random_digit</span></code>, which returns a value of any random integer between 0 and 9 (inclusive). (You can also see examples in the code.)</p>
<p>Write code that assigns to the variable <code class="docutils literal notranslate"><span class="pre">func_var</span></code> the <strong>function</strong> <code class="docutils literal notranslate"><span class="pre">greeting</span></code> (without executing the function).</p>
<p>Then, write code that assigns to the variable <code class="docutils literal notranslate"><span class="pre">new_digit</span></code> the <strong>return value</strong> from executing the function <code class="docutils literal notranslate"><span class="pre">random_digit</span></code>.</p>
<p>Then, write code that assigns to the variable <code class="docutils literal notranslate"><span class="pre">digit_func</span></code> the <strong>function</strong> <code class="docutils literal notranslate"><span class="pre">random_digit</span></code> (without executing the function).</p>

<textarea data-component="activecode" id=ps_1_05 data-lang="python" 
     data-include="addl_functions" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
# For example
print(greeting("Jackie"))
print(greeting("everybody"))
print(greeting("sdgadgsal"))

# Try running all this code more than once, so you can see how calling the function
# random_digit works.
print(random_digit())
print(random_digit())

# Write code that assigns the variables as mentioned in the instructions.


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(type(func_var), type(greeting), "Testing that func_var is same type as greeting")
    def testTwo(self):
       self.assertEqual(type(new_digit), type(1), "Testing that new_digit's value is an integer")
    def testThree(self):
       self.assertEqual(type(digit_func), type(random_digit), "Testing that digit_func is same type as random_digit")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ps_1_07" class="runestone explainer ac_section alert alert-warning">
<p>There is a function we are giving you for this problem set that takes two strings as inputs, and returns the length of both of those strings added together, called <code class="docutils literal notranslate"><span class="pre">add_lengths</span></code>. We are also including functions called <code class="docutils literal notranslate"><span class="pre">random_digit</span></code> and <code class="docutils literal notranslate"><span class="pre">square</span></code> in this problem set.</p>
<p>Now, take a look at the following code and related questions, in this code window.</p>

<textarea data-component="activecode" id=ps_1_07 data-lang="python" 
     data-include="addl_functions" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
new_str = "'Twas brillig"

y = add_lengths("receipt","receive")

x = random_digit()

z = new_str.find('b')

l = new_str.find("'")

# notice that this line of code is made up of a lot of different expressions
fin_value = square(len(new_str)) + (z - l) + (x * random_digit())

# DO NOT CHANGE ANY CODE ABOVE THIS LINE
# But below here, putting print statements and running the code may help you!

# The following questions are based on that code. All refer to the types of the
#variables and/or expressions after the above code is run.

#####################

# Write a comment explaining each of the following, after each question.
# Don't forget to press **run** to save once you've done this!

# What is square?

# What type of object does the expression square(len(new_str)) evaluate to?

# What type is z?

# What type is l?

# What type is the expression z-l?

# What type is x?

# What is random_digit? How many inputs does it take?

# What type does the expression x * random_digit() evaluate to?

# Given all this information, what type will fin_value hold once all this code is run?

====

print("===========")
print("\n\nThere are no tests for this problem")
</textarea>
</div>

<div data-childcomponent="ps_1_08" class="runestone explainer ac_section alert alert-warning">
<p>Assign the value of the third element of <code class="docutils literal notranslate"><span class="pre">num_lst</span></code> to a variable called <code class="docutils literal notranslate"><span class="pre">third_elem</span></code>.</p>
<p>Assign the value of the sixth element of <code class="docutils literal notranslate"><span class="pre">num_lst</span></code> to a variable called <code class="docutils literal notranslate"><span class="pre">elem_sixth</span></code>.</p>
<p>Assign the length of <code class="docutils literal notranslate"><span class="pre">num_lst</span></code> to a variable called <code class="docutils literal notranslate"><span class="pre">num_lst_len</span></code>.</p>
<p><em>Consider:</em> what is the difference between <code class="docutils literal notranslate"><span class="pre">mixed_bag[-1]</span></code> and <code class="docutils literal notranslate"><span class="pre">mixed_bag[-2]</span></code> (you may want to print out those values or print out information about those values, so you can make sure you know what they are!)?</p>
<p>Write code to print out the type of the third element of <code class="docutils literal notranslate"><span class="pre">mixed_bag</span></code>.</p>
<p>Write code to assign the <strong>type of the fifth element of</strong> <code class="docutils literal notranslate"><span class="pre">mixed_bag</span></code> to a variable called <code class="docutils literal notranslate"><span class="pre">fifth_type</span></code>.</p>
<p>Write code to assign the <strong>type of the first element of</strong> <code class="docutils literal notranslate"><span class="pre">mixed_bag</span></code> to a variable called <code class="docutils literal notranslate"><span class="pre">another_type</span></code>.</p>
<p><strong>Keep in mind:</strong> All ordinal numbers in <em>instructions</em>, like “third” or “fifth” refer to the way HUMANS count. How do you write code to find the right things?</p>

<textarea data-component="activecode" id=ps_1_08 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
num_lst = [4,16,25,9,100,12,13]
mixed_bag = ["hi", 4,6,8, 92.4, "see ya", "23", 23]

# Write your code here:


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(third_elem, num_lst[2], "Testing that third_elem has been set to the third element of num_lst")
    def testTwo(self):
       self.assertEqual(elem_sixth, num_lst[5], "Testing that elem_sixth has been set to the sixth element of num_lst")
    def testThree(self):
       self.assertEqual(num_lst_len,len(num_lst), "Testing that num_len has been set to the length of num_lst")
    def testFour(self):
       self.assertEqual(fifth_type, type(mixed_bag[4]), "Testing that fifth_type has been set to the type of the fifth element in mixed_bag")
    def testFive(self):
       self.assertEqual(another_type, type(mixed_bag[0]), "Testing that another_type has been set to the type of the first element of mixed_bag")
    def testSix(self):
       self.assertIn('print', self.getEditorText(), "Testing that 'print' is in your code. (Don't worry about Actual and Expected Values.)")
    def testSeven(self):
       self.assertIn('int', self.getOutput(), "Testing that you printed the correct element of mixed_bag. (Don't worry about Actual and Expected Values.)")


myTests().main()
</textarea>
</div>

        <div class="runestone" style="max-width: none;">
        <pre data-component="parsons" id="ps_1_09"      >
        <span data-question>Q-1:  Here's another complicated expression, using the Turtle framework we talked about. Arrange these sentences in the order they are executed in the following code, like you did in an exercise in Chapter 2 of the textbook. (It may help to think about what specifically is happening in the first four lines of code as well.)

 .. sourcecode:: python

      import turtle

      ella = turtle.Turtle()
      x = "hello class".find("o") - 1
      ella.speed = 3


      ella.move(square(x*ella.speed))

Order the code fragments in the order in which the Python interpreter would evaluate them, when evaluating that last line of code.

Not graded for pset points. But important practice!
</span>Look up the variable ella and find that it is an instance of a Turtle object
---
Look up the attribute move of the Turtle ella and find that it's a method object
---
Look up the function square
---
Look up the value of the variable x and find that it is an integer
---
Look up the value of the attribute speed of the instance ella and find that it is an integer
---
Evaluate the expression x * ella.speed to one integer
---
Call the function square on an integer value
---
Call the method .move of the Turtle ella on its input integer
        </pre>
        </div>
    
<div data-childcomponent="ps_1_10" class="runestone explainer ac_section alert alert-warning">
Write a program that uses the turtle module to draw something. It doesn’t have to be complicated, but draw something different than we did in the textbook or in class. (Optional but encouraged: post a screenshot of the artistic outcome to Piazza, or a short video of the drawing as it is created.) (Hint: if you are drawing something complicated, it could get tedious to watch it draw over and over. Try setting <code class="docutils literal notranslate"><span class="pre">.speed(10)</span></code> for the turtle to draw fast, or <code class="docutils literal notranslate"><span class="pre">.speed(0)</span></code> for it to draw super fast with no animation.)
<textarea data-component="activecode" id=ps_1_10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
import turtle
</textarea>
</div>
<p>That’s the end of the problem set. In the hidden code below, you will find the definitions of functions that were used elsewhere in the problem set. They’re hidden because you don’t yet need to understand how function definitions work. But if you want a preview, feel free to click on Show/Hide Code.</p>

<div data-childcomponent="addl_functions" class="runestone explainer ac_section alert alert-warning">
<strong>This is not a problem in the problem set. It is code that provides functions you’re using in this problem set. They’re hidden because you don’t yet need to understand how function definitions work.</strong>
<textarea data-component="activecode" id=addl_functions data-lang="python" 
    data-hidecode="true"  data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
def square(num):
    return num**2

def greeting(st):
    st = str(st) # just in case
    return "Hello, " + st

def random_digit():
    import random
    return random.choice([0,1,2,3,4,5,6,7,8,9])

def add_lengths(str1, str2):
    return len(str1) + len(str2)
</textarea>
</div>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Class Assignments' data-toggle="tooltip">
    <a href="../assignments.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="../assignments.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Activities: Week 2 (Note that PS1 is found in Week 1)' data-toggle="tooltip" >
    <a href="week02.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="week02.html" >Next Section - Activities: Week 2 (Note that PS1 is found in Week 1)</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span> | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2013 Paul Resnick, Brad Miller, David Ranum, Created using Runestone Interactive.
    </p>
  </div>
</footer>


<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>