<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Exercises &#8212; Programs, Information, and People: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/matrixeq.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webgldemo.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webglinteractive.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/js_lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../_static/accessibility.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="https://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Extra Exercises" href="ExtraExercises.html" />
    <link rel="prev" title="Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'pip2';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = true;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Table of Contents</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div class="gittip">
                  <iframe style="border:0;"
                    src="https://www.gittip.com/bnmnetp/widget.html"
                    width="48pt" height="22pt" class="gittip-button">
                  </iframe>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/pip2/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
<!--            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li> -->
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
<!--            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> --><!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/pip2/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='https://github.com/bnmnetp/runestone/issues/new'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Extra Exercises' data-toggle="tooltip" >
    <a href="ExtraExercises.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic simple">
<li>Write code to create a list of integers from 0 through 52 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. You should use a special Python function – do not type out the whole list yourself. HINT: You can do this in one line of code!</li>
</ol>

<div data-childcomponent="ee_ch10_01" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_01 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(numbers, range(53), "Testing that numbers is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>
<p>1.1 Write code to create a list of numbers from 0 to 67 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">nums</span></code>. Do not hard code the list.</p>

<div data-childcomponent="ee_ch10_011" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_011 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(nums, range(68), "Testing that nums is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="2">
<li>For each character in the string already saved in the variable <code class="docutils literal notranslate"><span class="pre">str1</span></code>, append each character to a list called <code class="docutils literal notranslate"><span class="pre">chars</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_02" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_02 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
str1 = "I love python"
# HINT: what's the accumulator? That should go here.

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(chars, ['I', ' ', 'l', 'o', 'v', 'e', ' ', 'p', 'y', 't', 'h', 'o', 'n'], "Testing that chars is assigned to correct values.")

myTests().main()
</textarea>
</div>
<p>2.1 For each character in the string saved in <code class="docutils literal notranslate"><span class="pre">ael</span></code>, append that character to a list that should be saved in a variable <code class="docutils literal notranslate"><span class="pre">app</span></code>.</p>

<div data-childcomponent="ee_ch10_022" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_022 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
ael = "python!"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(app, ['p','y','t','h','o','n', "!"], "Testing that app has the correct elements." )

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="3">
<li>Assign an empty string to the variable <code class="docutils literal notranslate"><span class="pre">output</span></code>. Using the <code class="docutils literal notranslate"><span class="pre">range</span></code> function, write code to make it so that the variable <code class="docutils literal notranslate"><span class="pre">output</span></code> has 35 <code class="docutils literal notranslate"><span class="pre">a</span></code> s inside it (like <code class="docutils literal notranslate"><span class="pre">&quot;aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa&quot;</span></code>). Hint: use the accumulation pattern!</li>
</ol>

<div data-childcomponent="ee_ch10_03" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_03 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(output, "aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa", "Testing that output has the correct value.")

myTests().main()
</textarea>
</div>
<p>3.1 Create an empty string and assign it to the variable <code class="docutils literal notranslate"><span class="pre">lett</span></code>. Then using range, write code such that when your code is run, lett has 7 b’s (<code class="docutils literal notranslate"><span class="pre">&quot;bbbbbbb&quot;</span></code>).</p>

<div data-childcomponent="ee_ch10_031" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_031 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(lett, "bbbbbbb", "Testing that lett has the correct value." )

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="4">
<li>Given the list of numbers, <code class="docutils literal notranslate"><span class="pre">numbs</span></code>, create a new list of those same numbers increased by 5. Save this new list to the variable <code class="docutils literal notranslate"><span class="pre">newlist</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_04" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_04 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
numbs = [5, 10, 15, 20, 25]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFour(self):
      self.assertEqual(newlist, [10, 15, 20, 25, 30], "Testing that the newlist value contains the correct elements.")

myTests().main()
</textarea>
</div>
<p>4.1 For each number in <code class="docutils literal notranslate"><span class="pre">lst_nums</span></code>, multiply that number by 2 and append it to a new list called <code class="docutils literal notranslate"><span class="pre">larger_nums</span></code>.</p>

<div data-childcomponent="ee_ch10_041" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_041 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
lst_nums = [4, 29, 5.3, 10, 2, 1817, 1967, 9, 31.32]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(larger_nums, [8, 58, 10.6, 20, 4, 3634, 3934, 18, 62.64], "Testing that larger_nums has been created correctly." )

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="5">
<li><strong>Challenge</strong> Now do the same as in problem 4, but do not create a new list. Overwrite the list <code class="docutils literal notranslate"><span class="pre">numbs</span></code> so that each of the original numbers are increased by 5.</li>
</ol>

<div data-childcomponent="ee_ch10_05" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_05 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
numbs = [5, 10, 15, 20, 25]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(numbs, [10, 15, 20, 25, 30], "Testing that numbs is assigned to correct values.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="6">
<li>For each word in the list <code class="docutils literal notranslate"><span class="pre">verbs</span></code>, add an -ing ending. Save this new list in a new list, <code class="docutils literal notranslate"><span class="pre">ing</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_06" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_06 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
verbs = ["kayak", "cry", "walk", "eat", "drink", "fly"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSix(self):
      self.assertEqual(ing, ['kayaking', 'crying', 'walking', 'eating', 'drinking', 'flying'], "Testing that the variable ing has the correct value.")

myTests().main()
</textarea>
</div>
<p>6.1 <strong>Challenge</strong> Do the same as above, but do not create a new list. Instead, overwrite the old list so that <code class="docutils literal notranslate"><span class="pre">verbs</span></code> has the same words with <code class="docutils literal notranslate"><span class="pre">ing</span></code> at the end of each one.</p>

<div data-childcomponent="ee_ch10_07" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_07 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
verbs = ["kayak", "cry", "walk", "eat", "drink", "fly"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSeven(self):
      self.assertEqual(verbs, ['kayaking', 'crying', 'walking', 'eating', 'drinking', 'flying'], "Testing that verbs is assigned to correct values.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="7">
<li>For each string in <code class="docutils literal notranslate"><span class="pre">wrds</span></code>, add ‘ed’ to the end of the word (to make the word past tense). Save these past tense words to a list called <code class="docutils literal notranslate"><span class="pre">past_wrds</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_061" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_061 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
wrds = ["end", 'work', "play", "start", "walk", "look", "open", "rain", "learn", "clean"]
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(past_wrds, ["ended", 'worked', "played", "started", "walked", "looked", "opened", "rained", "learned", "cleaned"], "Testing that past_wrds has the correct value." )

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="8">
<li>Count the number of characters in string <code class="docutils literal notranslate"><span class="pre">str1</span></code>. Do not use <code class="docutils literal notranslate"><span class="pre">len()</span></code>. Save the number in variable <code class="docutils literal notranslate"><span class="pre">numbs</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_08" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_08 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
str1 = "I like nonsense, it wakes up the brain cells. Fantasy is a necessary ingredient in living."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testEight(self):
      self.assertEqual(numbs, 90, "Testing that numbs is assigned to correct values.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="9">
<li>Create a list of numbers 0 through 40 and assign this list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. Then, accumulate the total of the list’s values and assign that sum to the variable <code class="docutils literal notranslate"><span class="pre">sum1</span></code>.</li>
</ol>

<div data-childcomponent="ee_ch10_09" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_ch10_09 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testNineA(self):
      self.assertEqual(numbers, [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40], "Testing that numbers is assigned to correct values.")

   def testNineB(self):
      self.assertEqual(sum1, 820, "Testing that sum1 has the correct value.")

myTests().main()
</textarea>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Glossary' data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Glossary.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Extra Exercises' data-toggle="tooltip" >
    <a href="ExtraExercises.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="ExtraExercises.html" >Next Section - Extra Exercises</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span> | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2013 Paul Resnick, Brad Miller, David Ranum, Created using Runestone Interactive.
    </p>
  </div>
</footer>


<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>