<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Exercises &#8212; Programs, Information, and People: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/matrixeq.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webgldemo.css" type="text/css" />
    <link rel="stylesheet" href="../_static/webglinteractive.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/js_lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../_static/accessibility.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="https://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Extra Exercises" href="ExtraExercises.html" />
    <link rel="prev" title=".CSV Output" href="CSV.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'https://506w18.learningpython.today' ? 'https://506w18.learningpython.today' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'pip2';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = true;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../index.html">Table of Contents</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div class="gittip">
                  <iframe style="border:0;"
                    src="https://www.gittip.com/bnmnetp/widget.html"
                    width="48pt" height="22pt" class="gittip-button">
                  </iframe>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/pip2/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
<!--            <li class="divider"></li>
            <li><a href='https://506w18.learningpython.today/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='https://506w18.learningpython.today/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='https://506w18.learningpython.today/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='https://506w18.learningpython.today/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li> -->
            <li><a href='https://506w18.learningpython.today/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='https://506w18.learningpython.today/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
<!--            <li><a href="https://506w18.learningpython.today/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="https://506w18.learningpython.today/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="https://506w18.learningpython.today/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> --><!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/pip2/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='https://github.com/bnmnetp/runestone/issues/new'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - .CSV Output" data-toggle="tooltip">
    <a href="CSV.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Extra Exercises' data-toggle="tooltip" >
    <a href="ExtraExercises.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<ol class="arabic simple">
<li>Using string interpolation, fill out the parameters so that “I go to UM and I am in SI 306” is assigned to <code class="docutils literal notranslate"><span class="pre">str1</span></code>.</li>
</ol>

<div data-childcomponent="ee_interpolation_01" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_01 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
str1 = "I go to {} and I am in {}".format()

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(str1, "I go to UM and I am in SI 306", "Testing that str1 is assigned to correct value")

myTests().main()
</textarea>
</div>
<p>1.1 Using string interpolation, fill out the parameters so that ‘This book teaches students python.’ is assigned to <code class="docutils literal notranslate"><span class="pre">info</span></code>.</p>

<div data-childcomponent="ee_interpolation_011" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_011 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
info = "This book {} students {}.".format()

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(info, "This book teaches students python.", "Testing that info has the correct value.")

myTests().main()
</textarea>
</div>
<p>1.2 Below, we have provided a list of words. Use string formatting to produce the following string: “My favorite animals are elephants, giraffes, and zebras.” Save this string to the variable name <code class="docutils literal notranslate"><span class="pre">animal_string</span></code>.</p>

<div data-childcomponent="ee_interpolation_012" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_012 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
animals = ['elephants', 'giraffes', 'zebras']

animal_string =

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(animal_string, "My favorite animals are elephants, giraffes, and zebras.", "Testing that animal_string is correct.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="2">
<li>Use string formatting to complete the string given. The blanks should correspond to variable <code class="docutils literal notranslate"><span class="pre">name</span></code> and <code class="docutils literal notranslate"><span class="pre">breed</span></code>.</li>
</ol>

<div data-childcomponent="ee_interpolation_02" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_02 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
name = "Oreo"
breed = "poodle"
str1 = "This is {}, he is a {}."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(str1, "This is Oreo, he is a poodle.", "Testing that str1 is assigned to correct value")

myTests().main()
</textarea>
</div>
<p>2.1 Using string interpolation, assign the correct value to the variable <code class="docutils literal notranslate"><span class="pre">names</span></code> so that the value assigned to the variable <code class="docutils literal notranslate"><span class="pre">sent</span></code> is “Paul, Jackie, and Stephen have taught or are teaching this class.”</p>

<div data-childcomponent="ee_interpolation_021" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_021 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
sent = "{}, {}, and {} have taught or are teaching this class.".format()

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(sent, "Paul, Jackie, and Stephen have taught or are teaching this class.", "Testing that sent has the correct value.")
      self.assertEqual(names, ['Paul', 'Jackie', 'Stephen'], "Testing that names has the correct values assigned")

myTests().main()
</textarea>
</div>
<p>2.2 Below, we have created the variables <code class="docutils literal notranslate"><span class="pre">course</span></code> and <code class="docutils literal notranslate"><span class="pre">school</span></code>. Use string formatting to produce the following string: “I’m enrolled in SI 106 here at University of Michigan.” Save this string to the variable name <code class="docutils literal notranslate"><span class="pre">final</span></code>.</p>

<div data-childcomponent="ee_interpolation_022" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_022 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
course = "SI 106"
school = "University of Michigan"

final =

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(final, "I'm enrolled in SI 106 here at University of Michigan.", "Testing that final is correct.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="3">
<li>Provided is a list of tuples, the first is a country, the second is their medal count. Create a new list called <code class="docutils literal notranslate"><span class="pre">medals</span></code> using these tuples so that if the tuple was (‘USA’, 121), then what is added to medals is the string “USA won 121 medals”. Do so using string interpolation.</li>
</ol>

<div data-childcomponent="ee_interpolation_03" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_03 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
countries = [('Jamaica', 11), ('Malaysia',5), ('Japan', 41), ('Sweden', 11), ('Serbia', 8)]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(medals, ['Jamaica won 11 medals', 'Malaysia won 5 medals', 'Japan won 41 medals', 'Sweden won 11 medals', 'Serbia won 8 medals'], "Testing that medals is assigned to correct values")

myTests().main()
</textarea>
</div>
<p>3.1 Provided is a list of tuples, the first is a name, the second is a city. Create a new list called <code class="docutils literal notranslate"><span class="pre">user_info</span></code> using these tuples so that if the tuple was (‘Ashley’, ‘Kalamazoo’), then what is added to user_info is the string “Ashley is from Kalamazoo”. Do so using string interpolation.</p>

<div data-childcomponent="ee_interpolation_031" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_031 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
info = [('Sarah', 'Mattawan'), ("Grace", "Kalamazoo"), ('Mariana', "Sao Paulo"), ('Kevin', 'Melbourne'), ('Srishti', 'Dubai'), ('Kathleen', 'Bagota'), ('Ann', 'Excel')]


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(user_info, ['Sarah is from Mattawan', 'Grace is from Kalamazoo', 'Mariana is from Sao Paulo', 'Kevin is from Melbourne', 'Srishti is from Dubai', 'Kathleen is from Bagota', 'Ann is from Excel'], "Testing that user_info has the correct value.")

myTests().main()
</textarea>
</div>
<p>3.2 Below, we have provided a list of tuples that contain information about summer Olympic meets. Create a new list called <code class="docutils literal notranslate"><span class="pre">olympics_info</span></code> using these tuples so that if the tuple is (‘2016’, ‘Rio de Janeiro, Brazil’), then what is added to <code class="docutils literal notranslate"><span class="pre">olympics_info</span></code> is the string: “The 2016 Olympics were held in Rio de Janeiro, Brazil.” Do this by using string interpolation.</p>

<div data-childcomponent="ee_interpolation_032" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_032 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
tups = [('2016', 'Rio de Janeiro, Brazil'), ('2012', 'London, Great Britain'), ('2008', 'Beijing, China'), ('2004', 'Athens, Greece'), ('2000', 'Sydney, Australia'), ('1996', 'Atlanta, Georgia, USA'), ('1992', 'Barcelona, Spain'), ('1988', 'Seoul, Korea')]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(olympics_info, ['The 2016 Olympics were held in Rio de Janeiro, Brazil.', 'The 2012 Olympics were held in London, Great Britain.', 'The 2008 Olympics were held in Beijing, China.', 'The 2004 Olympics were held in Athens, Greece.', 'The 2000 Olympics were held in Sydney, Australia.', 'The 1996 Olympics were held in Atlanta, Georgia, USA.', 'The 1992 Olympics were held in Barcelona, Spain.', 'The 1988 Olympics were held in Seoul, Korea.'], "Testing that olympics_info is correct.")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="4">
<li>Write a function called <code class="docutils literal notranslate"><span class="pre">pokemon</span></code> that takes in a list of an integer and string. The integer is the level of the trainer and the string is where the trainer plays. If the player is level five or below, they have the most rattatas. If they are between level 6 and 10, they have the most zubats. If they are higher than level 10, they have the most eevees. Return the string “I’m level __ and I caught a bunch of __ in the __!” where the first blank is the player level, the second is the pokemon, and the third is the location where they play. For instance, if the inputted list is [2, “city”], the returned string should be “I’m level 2 and I caught a bunch of rattatas in the city!” Do this using string interpolation.</li>
</ol>

<div data-childcomponent="ee_interpolation_04" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_04 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFour(self):
      self.assertEqual(pokemon([4, "suburbs"]), "I'm level 4 and I caught a bunch of rattatas in the suburbs!", "Testing that pokemon[4, 'suburbs'] returns 'I'm level 4 and I caught a bunch of rattatas in the suburbs!'.")
      self.assertEqual(pokemon([25, "field"]), "I'm level 25 and I caught a bunch of eevees in the field!", "Testing that pokemon[25, 'field'] returns 'I'm level 25 and I caught a bunch of eevees in the field!'.")
      self.assertEqual(pokemon([10, "city"]), "I'm level 10 and I caught a bunch of zubats in the city!", "Testing that pokemon[10, 'city'] returns 'I'm level 10 and I caught a bunch of zubats in the city!'.")

myTests().main()
</textarea>
</div>
<p>4.1 Write a function called <code class="docutils literal notranslate"><span class="pre">data_mine</span></code> that takes a tuple as input, and return a string using interpolation, depending on the second item in the tuple. The first item is a city name, the second is the weather condition. If the second item in the tuple has the value of ‘rain’, then it should also be cloudy. If the second item has the value of ‘sun’, then there should be blue skies. If the second item has the value of ‘snow’, then it should be cold. Return the string “In ___ there is ___ so it is ____.” or “In ____ there is ___ so there are _____.” (so for (‘Maz’, ‘rain’), we would return “In Maz there is rain so it is cloudy” (the same goes for snow), while for (‘Bieur’, ‘sun’), we would return “In Bieur there is sun so there are blue skies.”)</p>

<div data-childcomponent="ee_interpolation_041" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_041 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(data_mine(('Iron Mountain', 'snow')), "In Iron Mountain there is snow so it is cold.", "Testing that data_mine has the correct return value with input ('Iron Mountain', 'snow').")
      self.assertEqual(data_mine(('Santa Clara', 'sun')), 'In Santa Clara there is sun so there are blue skies.', "Testing that data_mine has the correct return value with input ('Santa Clara', 'sun')")
      self.assertEqual(data_mine(('Seattle', 'rain')), "In Seattle there is rain so it is cloudy.", "Testing that data_mine has the correct return value with input ('Seattle', 'rain')")

myTests().main()
</textarea>
</div>
<p>4.2 Write a function called <code class="docutils literal notranslate"><span class="pre">grades</span></code> that takes in a list with two elements, the first being a string (a person’s name) and the second being an integer (their grade on a test). If the grade is greater than or equal to 70, the function should return: “Congrats, [name], you passed the test with a [grade]!” If the grade is lower than 70, the function should return: “Sorry, [name], you failed the test with a [grade].”</p>

<div data-childcomponent="ee_interpolation_042" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_042 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(grades(["Jenny", 90]), "Congrats, Jenny, you passed the test with a 90!", "Testing the function grades on input ['Jenny', 90].")
      self.assertEqual(grades(["Tina", 70]), "Congrats, Tina, you passed the test with a 70!", "Testing the function grades on input ['Tina', 70].")
      self.assertEqual(grades(["Betty", 45]), "Sorry, Betty, you failed the test with a 45.", "Testing the function grades on input ['Betty', 45].")

myTests().main()
</textarea>
</div>
<ol class="arabic simple" start="5">
<li>The list of tuples, <code class="docutils literal notranslate"><span class="pre">order</span></code>, contains information about pizza orders. It contains information on whether or not the order is a pickup or delivery, how many pizzas were ordered, the kind of pizzas, and in how many minutes they need to be ready. Create a list called <code class="docutils literal notranslate"><span class="pre">response</span></code> that gives a response to each order. For a delivery, if the order input is (“delivery”, 1, “cheese”, 10), the response should be “Your 1 cheese pizza will be delivered in 10 minutes”. If the order is a pickup, the response should be “Come pick up your 1 cheese pizza in 10 minutes”.</li>
</ol>

<div data-childcomponent="ee_interpolation_05" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_05 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
order = [("delivery", 3, "pepperoni", 20), ("pickup", 4, "cheese", 10), ("pickup", 2, "combo", 5), ("delivery", 10, "cheese", 15), ("delivery", 1, "supreme", 60)]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(response, ['Your 3 pepperoni pizzas will be delivered in 20 minutes', 'Come pick up your 4 cheese pizzas in 10 minutes', 'Come pick up your 2 combo pizzas in 5 minutes', 'Your 10 cheese pizzas will be delievered in 15 minutes', 'Your 1 supreme pizzas will be delievered in 60 minutes'], "Testing if response is assigned to correct values")

myTests().main()
</textarea>
</div>
<p>5.1 Below, we have provided a list of tuples that contain information about customers’ product reviews on Amazon: the product, its rating, and customer name. Write a function called <code class="docutils literal notranslate"><span class="pre">feedback</span></code> that takes a tuple as input and returns a message to the customer based on their review. If the customer rated their product as an 8 or higher, <code class="docutils literal notranslate"><span class="pre">feedback</span></code> should return the following string: “[name], we’re happy to hear that you gave your new [product] a [rating] rating!” If the rating was below 8, <code class="docutils literal notranslate"><span class="pre">feedback</span></code> should return: “[name], we’re sorry to hear that your new [product] was not excellent.” Create a list called <code class="docutils literal notranslate"><span class="pre">feedback_messages</span></code> that contains a response to each customer below.</p>

<div data-childcomponent="ee_interpolation_052" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ee_interpolation_052 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
tups = [("Dyson vacuum", 9.1, "Sandy"), ("Keurig", 5.0, "Timmy"), ("SleepComfort mattress", 8.0, "Sam"), ("Michael Kors vest", 6.9, "Kate"), ("LG Dishwasher", 10.0, "Charles")]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(feedback_messages, ["Sandy, we're happy to hear that you gave your new Dyson vacuum a 9.1 rating!", "Timmy, we're sorry to hear that your new Keurig was not excellent.", "Sam, we're happy to hear that you gave your new SleepComfort mattress a 8 rating!", "Kate, we're sorry to hear that your new Michael Kors vest was not excellent.", "Charles, we're happy to hear that you gave your new LG Dishwasher a 10 rating!"], "Testing that feedback_messages is correct.")

myTests().main()
</textarea>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - .CSV Output' data-toggle="tooltip">
    <a href="CSV.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="CSV.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Extra Exercises' data-toggle="tooltip" >
    <a href="ExtraExercises.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="ExtraExercises.html" >Next Section - Extra Exercises</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span> | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2013 Paul Resnick, Brad Miller, David Ranum, Created using Runestone Interactive.
    </p>
  </div>
</footer>


<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>